#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Atomic Flag
// Test and set.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise atomic_flag not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
