#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Barrier
// Use std::barrier
// I AM NOT DONE

void solve() {
    // TODO: Use std::barrier.
    std::cout << "Exercise barrier not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
