#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Barrier Usage
// More complex barrier.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise barrier_usage not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
