#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Condition Variable
// Use condition_variable
// I AM NOT DONE

void solve() {
    // TODO: Wait for condition.
    std::cout << "Exercise condition not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
