#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Hardware Concurrency
// std::thread::hardware_concurrency.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise hardware_concurrency not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
