#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Lock Guard
// Use std::lock_guard
// I AM NOT DONE

void solve() {
    // TODO: Use std::lock_guard.
    std::cout << "Exercise lockguard not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
