#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Mutex
// Protect shared data with mutex
// I AM NOT DONE

void solve() {
    // TODO: Protect shared data.
    std::cout << "Exercise mutex1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
