#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Call Once
// Use std::call_once
// I AM NOT DONE

void solve() {
    // TODO: std::call_once.
    std::cout << "Exercise once not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
