#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Packaged Task
// std::packaged_task.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise packaged_task not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
