#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Recursive Mutex
// std::recursive_mutex.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise recursive_mutex not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
