#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Scoped Lock
// Use std::scoped_lock
// I AM NOT DONE

void solve() {
    // TODO: std::scoped_lock.
    std::cout << "Exercise scopedlock not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
