#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Semaphore
// std::counting_semaphore.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise semaphore not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
