#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Shared Mutex
// std::shared_mutex.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise shared_mutex not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
