#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Stop Token
// std::stop_token.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise stop_token not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
