#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Thread Local
// thread_local.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise thread_local_storage not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
