#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Thread Basics
// Create and join threads
// I AM NOT DONE

void solve() {
    // TODO: Start a thread.
    std::cout << "Exercise thread1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
