#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Byteswap
// Use std::byteswap
// I AM NOT DONE

void solve() {
    // TODO: std::byteswap.
    std::cout << "Exercise byteswap not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
