#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Cartesian Product
// views::cartesian_product.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise cartesian_product not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
