#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Deduce This
// Explicit object parameter
// I AM NOT DONE

void solve() {
    // TODO: Explicit object parameter.
    std::cout << "Exercise deducethis not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
