#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Expected
// Use std::expected
// I AM NOT DONE

void solve() {
    // TODO: Use std::expected.
    std::cout << "Exercise expected not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
