#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Flat Map
// Use std::flat_map
// I AM NOT DONE

void solve() {
    // TODO: Use std::flat_map.
    std::cout << "Exercise flatmap not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
