#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Forward Like
// Use std::forward_like
// I AM NOT DONE

void solve() {
    // TODO: std::forward_like.
    std::cout << "Exercise forward_like not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
