#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Join With
// views::join_with.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise join_with not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
