#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Move Only Function
// Use std::move_only_function
// I AM NOT DONE

void solve() {
    // TODO: std::move_only_function.
    std::cout << "Exercise moveonly not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
