#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Stacktrace
// std::stacktrace.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise stacktrace not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
