#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Extended Floats
// std::float16_t.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise stdfloat not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
