#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Unreachable
// Mark unreachable code
// I AM NOT DONE

void solve() {
    // TODO: Mark unreachable.
    std::cout << "Exercise unreachable not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
