#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Template Alias
// Use template aliases
// I AM NOT DONE

void solve() {
    // TODO: Using template.
    std::cout << "Exercise alias not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
