#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Deduction Guide
// User defined deduction guide.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise deduction_guide not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
