#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Default Template Args
// Default template parameters
// I AM NOT DONE

void solve() {
    // TODO: Default template args.
    std::cout << "Exercise default not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
