#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Dependent Name
// typename T::type.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise dependent_name not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
