#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Explicit Spec
// template<>.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise explicit_specialization not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
