#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Unary Fold
// (args + ...).
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise fold_unary not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
