#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Binary Fold
// (args + ... + init).
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise fold_binary not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
