#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Friend Templates
// Friend declarations in templates
// I AM NOT DONE

void solve() {
    // TODO: Friend declarations.
    std::cout << "Exercise friend not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
