#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Member Templates
// Template member functions
// I AM NOT DONE

void solve() {
    // TODO: Template inside class.
    std::cout << "Exercise member not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
