#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Non-type Parameters
// Use non-type template parameters
// I AM NOT DONE

void solve() {
    // TODO: Int template param.
    std::cout << "Exercise nontype not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
