#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Partial Specialization
// Partial template specialization
// I AM NOT DONE

void solve() {
    // TODO: Partial specialization.
    std::cout << "Exercise partial not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
