#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Template Recursion
// Recursive struct.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise recursion not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
