#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// SFINAE Class
// Class template SFINAE.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise sfinae_class not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
