#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Template Specialization
// Specialize templates
// I AM NOT DONE

void solve() {
    // TODO: Template specialization.
    std::cout << "Exercise specialization not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
