#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Template Function
// Create generic function
// I AM NOT DONE

void solve() {
    // TODO: Generic function.
    std::cout << "Exercise tempfunc not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
