#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Template Class
// Create generic class
// I AM NOT DONE

void solve() {
    // TODO: Generic class.
    std::cout << "Exercise tempclass not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
