#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Template Template Param
// template<template<class> class T>.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise template_template not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
