#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Variadic Templates
// Use parameter packs
// I AM NOT DONE

void solve() {
    // TODO: Variadic args.
    std::cout << "Exercise variadic not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
