#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Var Template Spec
// Specialization of variable.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise variable_template_spec not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
