#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Aligned Storage
// std::aligned_storage.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise aligned_storage not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
