#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Custom Allocator
// Create custom allocator
// I AM NOT DONE

void solve() {
    // TODO: Custom allocator basics.
    std::cout << "Exercise allocator not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
