#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Inline Assembly
// Use inline assembly
// I AM NOT DONE

void solve() {
    // TODO: Basic asm block.
    std::cout << "Exercise asm not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
