#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Assume Aligned
// std::assume_aligned.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise assume_aligned not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
