#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Attributes
// Use attributes
// I AM NOT DONE

void solve() {
    // TODO: [[nodiscard]], [[maybe_unused]].
    std::cout << "Exercise attributes not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
