#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Bit Fields
// Use bit fields
// I AM NOT DONE

void solve() {
    // TODO: Use bit fields.
    std::cout << "Exercise bitfields not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
