#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Construct At
// std::construct_at.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise construct_at not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
