#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Destroy At
// std::destroy_at.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise destroy_at not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
