#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Custom Exception
// Create custom exception
// I AM NOT DONE

void solve() {
    // TODO: Custom exception.
    std::cout << "Exercise exception not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
