#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Placement New
// Construct at specific address
// I AM NOT DONE

void solve() {
    // TODO: Construct at address.
    std::cout << "Exercise placementnew not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
