#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Polymorphic Memory Resource
// std::pmr.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise pmr not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
