#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// RTTI
// Runtime type information
// I AM NOT DONE

void solve() {
    // TODO: dynamic_cast and typeid.
    std::cout << "Exercise rtti not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
