#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// SIMD
// Placeholder for SIMD.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise simd_placeholder not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
