#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// To Address
// std::to_address.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise to_address not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
