#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Volatile
// Use volatile keyword
// I AM NOT DONE

void solve() {
    // TODO: Volatile keyword.
    std::cout << "Exercise volatile not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
