#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Common Type
// std::common_type.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise common_type not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
