#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Compile Time Math
// Factorial.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise compile_math not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
