#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Concepts for Metaprogramming
// Use concepts in metaprogramming
// I AM NOT DONE

void solve() {
    // TODO: Constraint based meta.
    std::cout << "Exercise concepts_meta not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
