#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Conditional
// Use std::conditional
// I AM NOT DONE

void solve() {
    // TODO: std::conditional.
    std::cout << "Exercise conditional not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
