#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Constexpr If
// Compile-time branching
// I AM NOT DONE

void solve() {
    // TODO: Compile time branch.
    std::cout << "Exercise constexpr_if not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
