#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// CRTP
// Curiously Recurring Template Pattern
// I AM NOT DONE

void solve() {
    // TODO: Curiously Recurring Template Pattern.
    std::cout << "Exercise crtp not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
