#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Enable If
// SFINAE with enable_if
// I AM NOT DONE

void solve() {
    // TODO: SFINAE with enable_if.
    std::cout << "Exercise enableif not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
