#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Integral Constant
// Use std::integral_constant
// I AM NOT DONE

void solve() {
    // TODO: std::integral_constant.
    std::cout << "Exercise integral not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
