#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Invoke Result
// std::invoke_result.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise invoke_result not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
