#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Is Detected
// Detection idiom.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise is_detected not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
