#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Policy Based Design
// Strategy as template param.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise policy not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
