#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Remove CV
// std::remove_cv.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise remove_cv not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
