#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Tag Dispatch
// Tag dispatching
// I AM NOT DONE

void solve() {
    // TODO: Dispatch based on tag.
    std::cout << "Exercise tagdispatch not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
