#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Type Traits
// Use type traits
// I AM NOT DONE

void solve() {
    // TODO: Use std::is_same.
    std::cout << "Exercise typetraits not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
