#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Type List
// Create type list
// I AM NOT DONE

void solve() {
    // TODO: Simple type list.
    std::cout << "Exercise typeliest not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
