#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Type List Ops
// Push front/back.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise typelist_ops not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
