#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Accumulate
// Sum elements
// I AM NOT DONE

void solve() {
    // TODO: Use std::accumulate.
    std::cout << "Exercise algo4 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
