#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Transform
// Use std::transform
// I AM NOT DONE

void solve() {
    // TODO: Transform elements.
    std::cout << "Exercise algo3 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
