#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Forward List
// Use std::forward_list.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise forward_list1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
