#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Iterators
// Use iterators
// I AM NOT DONE

void solve() {
    // TODO: Iterate with iterators.
    std::cout << "Exercise iter1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
