#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Multimap
// Use std::multimap.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise multimap1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
