#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Multiset
// Use std::multiset.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise multiset1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
