#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Priority Queue
// Use std::priority_queue.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise pqueue1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
