#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Strings
// Work with std::string
// I AM NOT DONE

void solve() {
    // TODO: Concatenate strings.
    std::cout << "Exercise string1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
