#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Unordered Map
// Use std::unordered_map.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise unordered_map1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
