#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Unordered Set
// Use std::unordered_set.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise unordered_set1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
