#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Vector Basics
// Create and use std::vector
// I AM NOT DONE

void solve() {
    // TODO: Create a vector and push elements.
    std::cout << "Exercise vec1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
