#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Vector Access
// Access vector elements
// I AM NOT DONE

void solve() {
    // TODO: Access elements safely.
    std::cout << "Exercise vec2 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
