#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Deserialize
// ParseFromString.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise proto4 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
