#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// File Input
// Read from file.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise proto18 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
