#!/usr/bin/env python3
"""
Welcome to PyTorchlings! 🔥

This is your first exercise. To complete it, simply remove the line
that says "I AM NOT DONE" below and run the exercise.

You can run this exercise with:
    python exercises/01_intro/intro01.py

Or use the pytorchlings runner:
    ./pytorchlings run intro01
"""

# I AM NOT DONE

print("Hello, PyTorch!")
print("Welcome to your PyTorch learning journey!")
print("Let's build some amazing deep learning models together!")

# Verify the exercise passed
assert True, "Exercise completed!"
