#!/usr/bin/env python3
"""
Let's verify PyTorch is installed correctly!

Your task: Import torch and print its version.
"""

# I AM NOT DONE

# TODO: Import torch
# import torch ???

# TODO: Print the PyTorch version
# print(f"PyTorch version: {???}")

# Verify torch is imported
import torch  # This import is for verification
assert torch is not None, "PyTorch imported successfully!"
print("✓ PyTorch is installed and working!")
