#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Auto Keyword
// Use auto for type deduction
// I AM NOT DONE

void solve() {
    // TODO: Use auto.
    std::cout << "Exercise auto1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
