#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Constexpr
// Use constexpr
// I AM NOT DONE

void solve() {
    // TODO: Make function constexpr.
    std::cout << "Exercise constexpr1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
