#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Default/Delete
// = default; = delete;
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise default_delete not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
