#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Delegating Constructors
// Call constructor from another.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise delegate not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
