#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Enum Class
// Use strongly typed enums
// I AM NOT DONE

void solve() {
    // TODO: Use enum class.
    std::cout << "Exercise enum1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
