#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Final Specifier
// Mark class/method final.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise final not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
