#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Inheriting Constructors
// using Base::Base.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise inherit not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
