#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Lambda Basics
// Create lambda functions
// I AM NOT DONE

void solve() {
    // TODO: Create a lambda.
    std::cout << "Exercise lambda1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
