#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Move Semantics
// Implement move constructor
// I AM NOT DONE

void solve() {
    // TODO: Implement move constructor.
    std::cout << "Exercise move1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
