#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Override Specifier
// Mark method override.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise override not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
