#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Range-based For
// Use range-based for loop
// I AM NOT DONE

void solve() {
    // TODO: Iterate over collection.
    std::cout << "Exercise rangefor not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
