#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Raw String Literals
// Use R"(...)".
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise rawstr not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
