#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Shared Pointer
// Use std::shared_ptr
// I AM NOT DONE

void solve() {
    // TODO: Use std::shared_ptr.
    std::cout << "Exercise smart2 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
