#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Unique Pointer
// Use std::unique_ptr
// I AM NOT DONE

void solve() {
    // TODO: Use std::unique_ptr.
    std::cout << "Exercise smart1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
