#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Trailing Return Type
// auto func() -> int.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise trailing not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
