#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Tuple
// Use std::tuple
// I AM NOT DONE

void solve() {
    // TODO: Create and unpack tuple.
    std::cout << "Exercise tuple1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
