#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// UTF-8 String Literals
// Use u8 prefix.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise u8str not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
