#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Concepts Basics
// Define a concept
// I AM NOT DONE

void solve() {
    // TODO: Define a simple concept.
    std::cout << "Exercise concepts1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
