#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Consteval
// Use immediate functions
// I AM NOT DONE

void solve() {
    // TODO: Immediate function.
    std::cout << "Exercise consteval not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
