#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Map Contains
// m.contains().
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise contains not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
