#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Coroutines
// Basic coroutine
// I AM NOT DONE

void solve() {
    // TODO: Basic coroutine handle.
    std::cout << "Exercise coroutines not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
