#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Designated Initializers
// {.x = 1, .y = 2}
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise designated not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
