#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// JThread
// Use std::jthread
// I AM NOT DONE

void solve() {
    // TODO: Use std::jthread.
    std::cout << "Exercise jthread not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
