#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Math Constants
// std::numbers::pi.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise numbers not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
