#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Ranges Sort
// Use ranges algorithms
// I AM NOT DONE

void solve() {
    // TODO: Use ranges::sort.
    std::cout << "Exercise ranges1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
