#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Ranges Views
// Use range views
// I AM NOT DONE

void solve() {
    // TODO: Use views::transform.
    std::cout << "Exercise ranges2 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
