#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Requires Clause
// Use requires clause
// I AM NOT DONE

void solve() {
    // TODO: Use requires.
    std::cout << "Exercise requires not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
