#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Source Location
// std::source_location.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise source_location not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
