#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Spaceship Operator
// Implement operator<=>
// I AM NOT DONE

void solve() {
    // TODO: Implement <=>.
    std::cout << "Exercise spaceship not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
