#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// String Starts With
// s.starts_with().
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise starts_with not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
