#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Template Lambda
// []<typename T>(T t){}
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise template_lambda not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
