#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Using Enum
// using enum Color.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise using_enum not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
