#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Adapter
// Implement Adapter pattern
// I AM NOT DONE

void solve() {
    // TODO: Implement Adapter.
    std::cout << "Exercise adapter not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
