#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Chain of Responsibility
// Implement Chain.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise chain not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
