#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Command
// Implement Command pattern
// I AM NOT DONE

void solve() {
    // TODO: Implement Command.
    std::cout << "Exercise command not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
