#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Composite
// Implement Composite pattern
// I AM NOT DONE

void solve() {
    // TODO: Implement Composite.
    std::cout << "Exercise composite not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
