#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Decorator
// Implement Decorator pattern
// I AM NOT DONE

void solve() {
    // TODO: Implement Decorator.
    std::cout << "Exercise decorator not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
