#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Factory Method
// Implement Factory pattern
// I AM NOT DONE

void solve() {
    // TODO: Implement Factory Method.
    std::cout << "Exercise factory not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
