#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Interpreter
// Implement Interpreter.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise interpreter not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
