#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Iterator Pattern
// Implement Iterator.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise iterator_pattern not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
