#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Observer
// Implement Observer pattern
// I AM NOT DONE

void solve() {
    // TODO: Implement Observer.
    std::cout << "Exercise observer not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
