#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Proxy
// Implement Proxy pattern
// I AM NOT DONE

void solve() {
    // TODO: Implement Proxy.
    std::cout << "Exercise proxy not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
