#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Singleton
// Implement Singleton pattern
// I AM NOT DONE

void solve() {
    // TODO: Implement Singleton.
    std::cout << "Exercise singleton not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
