#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Strategy
// Implement Strategy pattern
// I AM NOT DONE

void solve() {
    // TODO: Implement Strategy.
    std::cout << "Exercise strategy not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
