#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Template Method
// Implement Template Method.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise template_method not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
