#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Visitor
// Implement Visitor.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise visitor not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
