#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Reflection
// Server reflection.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise grpc18 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
