#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Keepalive
// Keepalive params.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise grpc20 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
