#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Server Builder
// Start server.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise grpc5 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
