#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Service Impl
// Implement service.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise grpc6 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
