#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Streaming Server
// Server streaming.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise grpc10 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
