#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Enum Attributes
// [[nodiscard]] enum class.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise attributes_enum not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
