#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Fallthrough
// [[fallthrough]]
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise fallthrough not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
