#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Filesystem
// Use std::filesystem
// I AM NOT DONE

void solve() {
    // TODO: Check file existence.
    std::cout << "Exercise filesystem not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
