#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Fold Expression
// Use fold expressions
// I AM NOT DONE

void solve() {
    // TODO: Use fold expression.
    std::cout << "Exercise fold not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
