#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// GCD LCM
// std::numeric::gcd.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise gcd_lcm not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
