#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// If with Initializer
// If statement with initializer
// I AM NOT DONE

void solve() {
    // TODO: If statement with initializer.
    std::cout << "Exercise ifinit not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
