#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Inline Variables
// inline static member.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise inline_var not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
