#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Make From Tuple
// std::make_from_tuple.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise make_from_tuple not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
