#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Maybe Unused
// [[maybe_unused]]
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise maybe_unused not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
