#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Nested Namespace
// namespace A::B {}
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise nested_namespace not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
