#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Nodiscard
// Mark return values
// I AM NOT DONE

void solve() {
    // TODO: Mark function nodiscard.
    std::cout << "Exercise nodiscard not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
