#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Optional
// Use std::optional
// I AM NOT DONE

void solve() {
    // TODO: Use std::optional.
    std::cout << "Exercise optional not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
