#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// String View Literals
// "sv"sv
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise string_view_literals not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
