#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// String View
// Use std::string_view
// I AM NOT DONE

void solve() {
    // TODO: Use std::string_view.
    std::cout << "Exercise stringview not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
