#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Structured Binding
// Unpack tuples/pairs
// I AM NOT DONE

void solve() {
    // TODO: Unpack pair/tuple.
    std::cout << "Exercise structbind not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
