#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Variant
// Use std::variant
// I AM NOT DONE

void solve() {
    // TODO: Use std::variant.
    std::cout << "Exercise variant not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
