#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Aggregate Initialization
// Struct init.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise aggr_init not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
