#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Aligned New
// operator new w/ alignment.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise aligned_new not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
