#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Binary Literals
// Use binary literals
// I AM NOT DONE

void solve() {
    // TODO: Use binary literals.
    std::cout << "Exercise binlit not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
