#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Deprecated with Msg
// [[deprecated("msg")]]
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise deprecated_msg not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
