#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Deprecated Attribute
// Mark functions deprecated
// I AM NOT DONE

void solve() {
    // TODO: Mark function deprecated.
    std::cout << "Exercise deprecate not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
