#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Digit Separators
// Use digit separators
// I AM NOT DONE

void solve() {
    // TODO: Use digit separators.
    std::cout << "Exercise digitsep not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
