#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Exchange
// Use std::exchange
// I AM NOT DONE

void solve() {
    // TODO: Use std::exchange.
    std::cout << "Exercise exchange not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
