#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Generic Lambda
// Use auto in lambda parameters
// I AM NOT DONE

void solve() {
    // TODO: Use auto in lambda args.
    std::cout << "Exercise genlambda not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
