#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Heterogeneous Lookup
// std::less<>.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise heterogeneous not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
