#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Integer Sequence
// std::integer_sequence.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise integer_sequence not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
