#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Lambda Capture Init
// [x = 1]
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise lambda_capture_init not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
