#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Make Unique
// Use std::make_unique
// I AM NOT DONE

void solve() {
    // TODO: Use std::make_unique.
    std::cout << "Exercise makeunique not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
