#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Relaxed Constexpr
// More complex constexpr
// I AM NOT DONE

void solve() {
    // TODO: More complex constexpr.
    std::cout << "Exercise relaxconst not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
