#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Auto Return Type
// Function auto return type deduction
// I AM NOT DONE

void solve() {
    // TODO: Function auto return.
    std::cout << "Exercise returntype not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
