#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Shared Lock
// shared_timed_mutex (C++14).
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise shared_lock not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
