#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Tuple Addressing
// std::get by type.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise tuple_address not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
