#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// User Defined Literals
// operator""_km.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise udl not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
