#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Variable Template
// Create variable templates
// I AM NOT DONE

void solve() {
    // TODO: Create variable template.
    std::cout << "Exercise varemplate not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
