#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Complex Variable Template
// Pi template with precision.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise vartmpl_complex not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
