#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Const
// Use const correctly.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise intro16 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
