#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Hello World
// Print 'Hello, World!'
// I AM NOT DONE

void solve() {
    // TODO: Make the code compile and print Hello, World!
    std::cout << "Exercise intro1 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
