#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Do While Loop
// Use do-while.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise intro12 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
