#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Switch Statement
// Use switch/case.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise intro11 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
