#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Basic Math
// Perform arithmetic
// I AM NOT DONE

void solve() {
    // TODO: Calculate sum of a and b.
    std::cout << "Exercise intro3 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
