#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Command Line Args
// Read argc/argv.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise intro20 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
