#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// While Loop
// Use while loop
// I AM NOT DONE

void solve() {
    // TODO: Loop until condition met.
    std::cout << "Exercise intro6 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
