#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// If Statement
// Use conditional logic
// I AM NOT DONE

void solve() {
    // TODO: Fix the condition.
    std::cout << "Exercise intro4 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
