#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Global Variables
// Use extern or avoid global.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise intro15 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
