#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Static Cast
// Use static_cast.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise intro17 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
