#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Function Parameters
// Pass arguments to functions
// I AM NOT DONE

void solve() {
    // TODO: Pass arguments.
    std::cout << "Exercise intro8 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
