#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Return Values
// Return from functions
// I AM NOT DONE

void solve() {
    // TODO: Return the correct value.
    std::cout << "Exercise intro9 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
