#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Functions
// Define and call functions
// I AM NOT DONE

void solve() {
    // TODO: Call the function.
    std::cout << "Exercise intro7 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
