#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// References
// Use references
// I AM NOT DONE

void solve() {
    // TODO: Use a reference to modify variable.
    std::cout << "Exercise intro10 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
