#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// Variables
// Declare and initialize variables
// I AM NOT DONE

void solve() {
    // TODO: Initialize the variable 'x' to 10.
    std::cout << "Exercise intro2 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
