#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// C Strings
// Use char* and strlen.
// I AM NOT DONE

void solve() {
    // TODO: implementation
    std::cout << "Exercise intro19 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
