#include <iostream>
#include <vector>
#include <string>
#include <algorithm>

// For Loop
// Iterate with for loop
// I AM NOT DONE

void solve() {
    // TODO: Iterate 10 times.
    std::cout << "Exercise intro5 not implemented!" << std::endl;
    // exit(1);
}

int main() {
    solve();
    return 0;
}
